1. Write a function that called **squareroot** that returns the squareroot of a number

<!-- grade_id: cell-2d2f2253c6076055, points: 0, grade: false, solution: true -->
```python
def squareroot(number):
     return number**0.5
print(squareroot(100))
print(squareroot(50))

```

Output:
```
10.0
7.0710678118654755
```

<!-- grade_id: cell-e64a12da4fca8b8d, points: 2, grade: true, solution: false -->
```python
assert squareroot(100) == 10
assert squareroot(25) == 5
assert squareroot(4) == 2
```

2. Create a function to calculate the discounted price on a product given the original price and the percentage discount.
The formula for calculating discount is
$$
\text{discount} = \left(\frac{\text{discount\_percentage}}{100}\right) * \text{original\_price}
$$

$$
\text{discounted\_price} = \text{original\_price} - \text{discount}
$$

<!-- grade_id: cell-ece99dbbbb868f82, points: 0, grade: false, solution: true -->
```python
def calculate_discount(original_price, discount_percentage):
    #calculte discount
    result = original_price * discount_percentage /100 
    discounted_price = original_price - result
    return discounted_price

calculate_discount(500,10)

```

Output:
```
450.0
```

<!-- grade_id: cell-756ff72b038890c7, points: 2, grade: true, solution: false -->
```python
assert calculate_discount(500, 10) == 450.0
assert calculate_discount(100, 20) == 80.0
assert calculate_discount(1500, 15) == 1275.0
```

3. using the same formula in question 2, write a function called **discount_calculator** that will calculate the discounted price of a product using a default discount of 5% unless a different percentage is specified.

<!-- grade_id: cell-6e800e3609ad0dc4, points: 0, grade: false, solution: true -->
```python
# writing a function that will calculate a default discount
def discount_calculator(original_price, percentage_discount=5):
    result = original_price * percentage_discount/100
    discounted_price = original_price -result
    if percentage_discount == 5:
        return discounted_price
    
    else:
        result = original_price * percentage_discount /100
        discounted_price = original_price -result
        return discounted_price
    
discount_calculator(100)
```

Output:
```
95.0
```

<!-- grade_id: cell-0a6c826e8a5ac374, points: 2, grade: true, solution: false -->
```python
assert discount_calculator(500, 10) == 450.0
assert discount_calculator(100) == 95.0
assert discount_calculator(1500, 15) == 1275.0
assert discount_calculator(1000) == 950.0
```

4. Write a function called **grade** that will return the grade of student depending on the score.
   The scoring template is as follows:
   
   * Above 70 - A
   * 60 to 69 - B
   * 50 to 59 - C
   * 45 to 49 - D
   * 40 to 44 - E
   * Below 40 - F

<!-- grade_id: cell-233ab09ca75b996d, points: 0, grade: false, solution: true -->
```python
#creating a grade function to return student grades
def grade(score):
    if score >= 70:
        grade = "A"
    elif score >= 60:
          grade = "B"
    elif score >= 50:
          grade = "C"
    elif score >= 45:
          grade = "D"
    elif score >= 40:
          grade = "E"
    else:
         grade = "F"
    return grade
grade(42)
```

Output:
```
'E'
```

<!-- grade_id: cell-093011897cd98f6b, points: 2, grade: true, solution: false -->
```python
assert grade(90) == 'A'
assert grade(33) == 'F'
assert grade(65) == 'B'
assert grade(47) == 'D'
```

5. Write a function called **highest** that returns the highest value in a list.

<!-- grade_id: cell-82c763f423b0be92, points: 0, grade: false, solution: true -->
```python
# writing a function to return the highest value in a list
def highest(list1):
    c = list1[0]
    for item in list1:
        if item > c:
            c =item
    return c
highest([30,100,600])
```

Output:
```
600
```

<!-- grade_id: cell-7f19173d420da15d, points: 2, grade: true, solution: false -->
```python
assert highest([10, 20, 3.5]) == 20
assert highest(['A', 'B', 'Z', 'K']) == 'Z'
```